# Cours ENPC - Pratique du calcul scientifique

Before you turn in this assignment, make sure everything runs as expected. First, **restart the kernel** and then **run all cells**. Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and group number below:

```julia
NAME = ""
GROUP = ""
```

## Week 3: Linear systems

```julia
using LaTeXStrings
using LinearAlgebra
using Plots
using Polynomials
```

### <font color='green'>Conditioning</font>

### <font color='orange'>[Exercise 1]</font> Link between conditioning and relative error

Consider the linear system
$$
    \mathsf A(α) \mathbf x :=
    \begin{pmatrix}
        1 & 1 \\
        1 & 1 - α
    \end{pmatrix}
    \begin{pmatrix}
        x_1 \\
        x_2
    \end{pmatrix} =
    \begin{pmatrix}
        π \\
        π - πα
    \end{pmatrix} =: \mathbf b(α).
$$
For $α$ in the range $[10^{-15}, 10^{-2}]$,
plot on the same graph, using logarithmic scales for both axes,

- The function which to $α$ associate $\kappa_2\bigl(\mathsf A(α)\bigr) \times ε$,
  where $\kappa_2\bigl(\mathsf A(α)\bigr)$ is the condition number of the matrix $A(α)$ for the Euclidean norm,
  and $ε$ is the machine epsilon for the `Float64` format.

- The relative error, in Euclidean norm, obtained when solving the linear system with the backslash `\` operator,
  in the case where both $A(α)$ and $b(α)$ are stored in the `Float64` format.

<!-- grade_id: conditining_plot, points: 2, grade: true, solution: true -->
```julia
Af(α) = [1. 1.; 1. (1-α)]
bf(α) = [π; π - π*α]

# Exact solution
x_exact = [0.; π]

# Range of α
αs = 10. .^ (-15:.1:-2)

# YOUR CODE HERE
throw(ErrorException("No code provided"))
```

### <font color='green'>Direct methods</font>

### <font color='orange'>[Exercise 2]</font> LU decomposition without Gaussian elimination

L'objectif de cet exercice est de proposer un algorithme permettant de réaliser la décomposition LU d'une matrice réelle $\mathsf{A}\in\mathbb{R}^{n×n}$,
**non pas par élimination gaussienne** mais par identification des entrées de $\mathsf A$ avec celles de $\mathsf L \mathsf U$.
Il s'agit de trouver un matrice triangulaire inférieure $\mathsf L$ formée de 1 sur la diagonale
et une matrice triangulaire supérieure $\mathsf U$ telles que :
<a id="LU"></a>
$$
\tag{LU}
\mathsf{A}=\mathsf{L}\mathsf{U}
$$

1. Écrire une fonction `my_lu(A)` qui prend comme argument une matrice `A` et qui renvoie les matrices `L` et `U`.
   Pour calculer ces matrices, s'appuyer sur une identification successive des éléments des deux membres de <a href="#LU">(LU)</a>,
   ligne par ligne de haut en bas, et de gauche à droite au sein de chaque ligne.

   <details>
       <summary>
           <em><font color='gray'>Hint (click to display)</font></em>
       </summary>

   Lorsqu'on suit l'ordre conseillé,
   la comparaison de l'élément $(i, j)$ fournit une équation pour $\ell_{ij}$ si $j < i$,
   et une équation pour $u_{ij}$ si $j \geq i$.
   Notons qu'il est possible de parcourir les éléments dans d'autres ordres.
   </details>

<!-- grade_id: my_lu, points: 0, grade: false, solution: true -->
```julia
function my_lu(A)
    # YOUR CODE HERE
    throw(ErrorException("No code provided"))
    return L, U
end;
```

<!-- grade_id: my_lu_tests, points: 3, grade: true, solution: false -->
```julia
ref_lu(A) = LinearAlgebra.lu(A, NoPivot())
@assert my_lu(diagm([1; 2; 3])) == (diagm([1; 1; 1]), diagm([1; 2; 3]))
@assert my_lu([2 -1 0; -1 2 -1; 0 -1 2])[1] ≈ [1 0 0; -1/2 1 0; 0 -2/3 1]
@assert my_lu([2 -1 0; -1 2 -1; 0 -1 2])[2] ≈ [2 -1 0; 0 3/2 -1; 0 0 4/3]
@assert (C = [1 2 3 4; 4 3 2 1; 1 2 1 2; 1 5 4 1]; my_lu(C)[1] ≈ ref_lu(C).L)
@assert (C = [1 2 3 4; 4 3 2 1; 1 2 1 2; 1 5 4 1]; my_lu(C)[2] ≈ ref_lu(C).U)
@assert (C = randn(100, 100); my_lu(C)[1] ≈ ref_lu(C).L)
@assert (C = randn(100, 100); my_lu(C)[2] ≈ ref_lu(C).U)
```

2. On suppose maintenant que la matrice réelle définie positive `A` est à largeur de bande `b` supposée beaucoup plus petite que `n`.
   Réécrire la fonction de décomposition LU en exploitant la largeur de bande.

   <details>
       <summary>
           <em><font color='gray'>Hint (click to display)</font></em>
       </summary>

   Pour rappel, la largeur de bande est le plus petit nombre naturel $b$ tel que $a_{ij} = 0$ for all $i,j \in \{1, \dots, n\}^2$ such that $|i-j| > b$.
   </details>

<!-- grade_id: my_banded_lu, points: 0, grade: false, solution: true -->
```julia
function my_banded_lu(A, b)
    # YOUR CODE HERE
    throw(ErrorException("No code provided"))
    return L, U
end;
```

<!-- grade_id: cell-24fb0131261b7be2, points: 2, grade: true, solution: false -->
```julia
@assert begin C = randn(100, 100); my_banded_lu(C, 99)[1] ≈ ref_lu(C).L end
@assert begin C = randn(100, 100); my_banded_lu(C, 99)[2] ≈ ref_lu(C).U end

```

3. Construire une fonction `generate_banded(n, b)` permettant de générer une matrice carrée aléatoire de taille `n` à largeur de bande donnée `b`.

<!-- grade_id: cell-30b31cd2999ed291, points: 0, grade: false, solution: true -->
```julia
function generate_banded(n, b)
    # YOUR CODE HERE
    throw(ErrorException("No code provided"))
end;
```

<!-- grade_id: generate_banded, points: 2, grade: true, solution: false -->
```julia
@assert generate_banded(10, 2)[1, 5] == 0
@assert generate_banded(10, 2)[2, 5] == 0
@assert generate_banded(10, 2)[3, 5] != 0
@assert generate_banded(10, 2)[4, 5] != 0
@assert generate_banded(10, 2)[5, 5] != 0
@assert generate_banded(10, 2)[6, 5] != 0
@assert generate_banded(10, 2)[7, 5] != 0
@assert generate_banded(10, 2)[8, 5] == 0
@assert generate_banded(10, 2)[9, 5] == 0
```

4. En utilisant `generate_banded`, tester votre implémentation de `my_banded_lu`,
   pour `n = 100` et des valeurs de `b` égales à 2, 3 et 10.
   <details>
       <summary>
           <em><font color='gray'>Hint (click to display)</font></em>
       </summary>

   Vous pouvez utiliser la fonction `lu` de la bibliothèque `LinearAlgebra`,
   avec l'argument `NoPivot()`, correspondant à la fonnction `ref_lu` définie ci-dessus,
   comme fonction de référence.
   Vous pouvez également utiliser la macro `@assert` pour vos tests.
   </details>

<!-- grade_id: cell-9dd7b3831c6e1227, points: 1, grade: true, solution: true -->
```julia
# YOUR CODE HERE
throw(ErrorException("No code provided"))
```

### <font color='green'>Iterative methods</font>

### <font color='orange'>[Exercise 3]</font> Richardson's iteration

Considérer le système linéaire suivant:
$$
    \mathsf A \mathbf x :=
    \begin{pmatrix}
        3 & 1 \\ 1 & 3
    \end{pmatrix}
    \begin{pmatrix}
        x_1 \\
        x_2
    \end{pmatrix}
    =
    \begin{pmatrix}
        11 \\
        9
    \end{pmatrix} =: \mathbf b.
$$

 1. Illustrer à l'aide de la fonction `Plots.contourf` les lignes de niveau de la fonction
 $$
     f(\mathbf x) = \frac{1}{2} \mathbf x^T \mathsf A \mathbf x - \mathbf b^T \mathbf x.
 $$

<!-- grade_id: richardson, points: 1, grade: true, solution: true -->
```julia
A = [3. 1.; 1. 3.]
b = [11.; 9.]
sol = A\b

# YOUR CODE HERE
throw(ErrorException("No code provided"))
```

 2. Implémenter l'itération de Richardson avec $\omega = 0.1$ pour résoudre le système.
    Votre fonction devra renvoyer une liste contenant toutes les itérations.
    Initialiser l'algorithme à $\mathbf x = 0$ et,
    comme critère d'arrêt, utiliser
    $$
    \lVert \mathsf A \mathbf x - \mathbf b \lVert \leq \varepsilon \lVert \mathbf b \lVert, \qquad \varepsilon = 10^{-50}
    $$
    <details>
        <summary>
            <em><font color='gray'>Hint (click to display)</font></em>
        </summary>

    Pour ajouter un élément à la fin d'une liste,
    utiliser la fonction `push!`:

    ```julia
    push!(xs, x)  # Adds x at the end of xs
    ```
    </details>

<!-- grade_id: cell-2304e66a1cccd8ea, points: 2, grade: true, solution: true -->
```julia
function richardson(ω)
    ε = 1e-50
    x = zeros(BigFloat, 2)
    xs = [x]
    # YOUR CODE HERE
    throw(ErrorException("No code provided"))
    return xs
end

ω = .1
xs = richardson(ω)
scatter!(eachrow(hcat(xs...))...)
plot!(eachrow(hcat(xs...))...)
```

 3. Faire un plot de la norme du résidu $\lVert r_k \rVert := \lVert \mathsf A \mathbf x^{(k)} - \mathbf b \rVert$ en fonction de $k$,
    en utilisant une échelle linéaire pour l'axe des abcisses et une échelle logarithmique pour l'axe des ordonnées,
    gràce à l'argument `yscale=:log` passé à la fonction `Plots.plot`.

<!-- grade_id: plot_error, points: 1, grade: true, solution: true -->
```julia
# YOUR CODE HERE
throw(ErrorException("No code provided"))
```

 4. En utilisant `Polynomials.fit`, calculer une approximation du type
    $$
    \log(e_k) \approx α + βk \qquad \Leftrightarrow \qquad e_k \approx \exp(α) \times \exp(β)^k.
    $$
    Comparer la valeur de $\exp(β)$ au rayon spectral $\rho(\mathsf A - \omega \mathsf I)$.

<!-- grade_id: cell-0dcbe1c4e4e61806, points: 1, grade: true, solution: true -->
```julia
# Define b and ρ below
# YOUR CODE HERE
throw(ErrorException("No code provided"))

exp(β), ρ
```

5. Calculer le paramètre $\omega$ optimal et refaire le plot de la décroissance de la norme du résidu dans ce cas.

<!-- grade_id: cell-f2d06f69cd09a3aa, points: 1, grade: true, solution: true -->
```julia
# YOUR CODE HERE
throw(ErrorException("No code provided"))
```

### <font color='orange'>[Exercise 4]</font> Basic iterative method for the Euler-Bernoulli beam

The objective of this exercise is to solve the Euler-Bernoulli beam equation in one dimension,
with homogeneous Dirichlet boundary conditions:
$$
u''''(x) = 1, \qquad u(0) = u'(0) = u'(1) = u(1) = 0.
$$
This equation models the deflection of a beam clamped at both ends,
under a uniform load.
A discretization of this equation on a uniform grid using the finite difference method leads to the following linear system:
$$
\begin{pmatrix}
    6  & -4 & 1 \\
    -4 & 6  & -4 & 1 \\
     1 & -4 & 6      & -4 & 1 \\
     & 1 & -4 & 6      & -4 & 1 \\
     &  &  \ddots  & \ddots & \ddots & \ddots & \ddots  \\
     & &    &   1    & -4    & 6      & -4 & 1 \\
     & & &    &   1    & -4    & 6      & -4 \\
     & & &    &        &  1  & -4      & 6 \\
\end{pmatrix}
\begin{pmatrix}
    u_2 \\
    u_3 \\
    u_4 \\
    u_5 \\
    \vdots \\
    u_{n-4} \\
    u_{n-3} \\
    u_{n-2}
\end{pmatrix}
=
h^4
\begin{pmatrix}
    1 \\
    1 \\
    1 \\
    1 \\
    \vdots \\
    1 \\
    1 \\
    1
\end{pmatrix},
\quad
h := \frac{1}{n},
\quad
x_i := ih.
$$
where $h$ is the spacing between the discretization points, and $(u_2, u_3, \cdots, u_{n-3}, u_{n-2})$ are the values of the unknown function $u$ at the points $(x_2, x_3, \cdots, x_{n-3}, x_{n-2})$.

1. Write a function `build_matrix(n)`, which returns the matrix of the linear system.

   <details>
       <summary>
           <em><font color='gray'>Hint (click to display)</font></em>
       </summary>

   The function `LinearAlgebra.diagm` is useful to construct the matrix.
   For example, the command
   ```julia
   diagm(0 => [1, 2, 3], -1 => [5, 5])  # = [1 0 0; 5 2 0; 0 5 3]
   ```
   returns a 3x3 matrix with `[1, 2, 3]` on the main diagonal,
   and `[5, 5]` on the first subdiagonal.
   </details>

<!-- grade_id: cell-71834b89fc893e7b, points: 0, grade: false, solution: true -->
```julia
function build_matrix(n)
    # YOUR CODE HERE
    throw(ErrorException("No code provided"))
end
```

<!-- grade_id: cell-ef8e08b0ec3d4d5b, points: 1, grade: true, solution: false -->
```julia
@assert size(build_matrix(20)) == (17, 17)
@assert build_matrix(20)[5, 3] ≈ 1
@assert build_matrix(20)[5, 4] ≈ -4
@assert build_matrix(20)[5, 5] ≈ 6
@assert build_matrix(20)[5, 6] ≈ -4
@assert build_matrix(20)[5, 7] ≈ 1
@assert build_matrix(20)[5, 8] ≈ 0
```

2. Write a function `build_rhs(n)`, which returns the right-hand side of the linear system.

<!-- grade_id: cell-ae67b79c0f349db8, points: 0, grade: false, solution: true -->
```julia
function build_rhs(n)
    # YOUR CODE HERE
    throw(ErrorException("No code provided"))
end
```

<!-- grade_id: cell-442b10910f62252e, points: 1, grade: true, solution: false -->
```julia
@assert length(build_rhs(20)) == 17
@assert build_rhs(20)[end] == 1 / 20^4
@assert build_rhs(20)[1] == 1 / 20^4
```

3. Write a function `forward_substitution(L, y)`
   that solves the linear system $\mathsf L \mathbf x = \mathbf y$,
   in the case where `L` is a lower-triangular matrix,
   by using forward substitution.

<!-- grade_id: cell-8abbc441ef2de365, points: 0, grade: false, solution: true -->
```julia
function forward_substitution(L, y)
    # YOUR CODE HERE
    throw(ErrorException("No code provided"))
end
```

<!-- grade_id: cell-4085a1598f1048e2, points: 1, grade: true, solution: false -->
```julia
@assert begin n = 10; L = [1 0; 2 1]; forward_substitution(L, [1; 0]) ≈ [1; -2] end
@assert begin n = 10; L = LowerTriangular(ones(n, n)); y = fill(2, n); forward_substitution(L, L*y) ≈ y end
@assert begin n = 10; L = LowerTriangular(randn(n, n)); y = randn(n); forward_substitution(L, L*y) ≈ y end
@assert begin n = 20; L = LowerTriangular(2ones(n, n)); y = rand(n); forward_substitution(L, L*y) ≈ y end
```

4. The successive over-relaxation method is a splitting method for solving linear systems of the form $\mathsf A \mathbf x = \mathbf b$.
   It is based on the iteration
   $$
   \mathsf M \mathbf x_{k+1} = \mathsf N \mathbf x_{k} + \mathbf{b}, \qquad \text{where} \quad \mathsf M = \frac{\mathsf D}{\omega} + \mathsf L \quad \text{and} \quad \mathsf N = \mathsf M - \mathsf A.
   \tag{SOR}
   $$
   <a id="SOR"></a>
   Here $\omega \in (0, 2)$ is a parameter,
   $\mathsf D$ is diagonal matrix containing the diagonal of $\mathsf A$,
   and $\mathsf L$ is a strictly lower triangular matrix containing the strict lower triangular part of $\mathsf A$,
   not including the diagonal.
   Write a function `sor(A, b, ω)` that
   implements this iteration, without using Julia's `\` and `inv` functions.
   Initialize the iteration at $\mathbf x_0 = \mathbf 0$,
   and use as stopping criterion that
   $$
   \lVert \mathsf A \mathbf x - \mathbf b \rVert
   \leq \varepsilon \lVert \mathbf b \rVert,
   \qquad \varepsilon := 10^{-10}.
   $$
   If after `maxiter` iterations,
   a solution that satisfies this stopping criterion has not been found,
   return `nothing`.

   <details>
       <summary>
           <em><font color='gray'>Hint (click to display)</font></em>
       </summary>

   - The functions `Diagonal` and `LowerTriangular`,
     both from the `LinearAlgebra` library,
     are useful to construct the matrices $\mathsf D$ and $\mathsf L$:
     ```julia
     D = Diagonal(A)  # Extracts the diagonal from A
     X = LowerTriangular(A)  # Extracts lower triangular part, with diag
     ```
   - Since the matrix in the linear system <a href="#NR">(SOR)</a> is lower triangular,
     use the function `forward_substitution` you wrote above to solve this system at each iteration.
   </details>

<!-- grade_id: cell-048154d0c43ed3fd, points: 0, grade: false, solution: true -->
```julia
function sor(A, b; ω = 1, ε = 1e-10, maxiter = 10^5)
    # YOUR CODE HERE
    throw(ErrorException("No code provided"))
end
```

<!-- grade_id: cell-8af811797a39d253, points: 2, grade: true, solution: false -->
```julia
# Test code
n = 20
X = qr(randn(n, n)).Q
A = X*Diagonal(rand(n) .+ 5)*X'
b = randn(n)
@assert begin ε = 1e-10; sol = sor(A, b; ω = 1.5, ε = ε); norm(A*sol - b) ≤ ε*norm(b) end
@assert begin ε = 1e-10; sol = sor(A, b; ω = 1.5, ε = ε); norm(A*sol - b) ≥ 1e-15*norm(b) end
@assert begin ε = 1e-12; sol = sor(A, b; ω = 1.5, ε = ε); norm(A*sol - b) ≤ ε*norm(b) end
@assert begin ε = 1e-12; sor(A, b; ω = 2, ε = ε) == nothing end
@assert begin ε = 1e-12; sor(A, b; ω = 1, ε = ε) ≈ A\b end
```

**Remark**. Here we used a function with both non-keyword (before `;`) and keyword (after `;`) arguments,
as well as default values for the keyword arguments.
When default values are present, the arguments become optional.
The following example should help you understand how arguments are handled in Julia.

```julia
my_test(a , b = 2; c = 3, d = 4) = "Arguments: $a $b $c $d"
@show my_test(5)
@show my_test(5, 8)
@show my_test(5, 8; c = 9)
@show my_test(5, 8; d = 9)
@show my_test(5; d = 9);

# my_test(5, 8, 9)  # -> Error because there are only two non-keyword arguments
# my_test(5, 8; 9)  # -> Error because there should only be keyword arguments after ';'
```

5. Use the relaxation method implemented in the previous item
   with parameters $\omega = 1.9$ and $\varepsilon = 10^{-8}$ to solve the linear system of this exercise with $n = 50$.
   Compare on a graph the solution you obatined with the exact solution,
   which is given by
   $$ u(x) = \frac{1}{24} x^2(x - 1)^2. $$

<!-- grade_id: comparison, points: 1, grade: true, solution: true -->
```julia
u(x) = 1/24 * x^2 * (x - 1)^2

# YOUR CODE HERE
throw(ErrorException("No code provided"))
```